<problem url_name="Custom_Python_Evaluated_Input_0" display_name="Custom Python-Evaluated Input" markdown="null">
  <p>
A custom python-evaluated input problem accepts one or more lines of text input from the student, and evaluates the inputs for correctness based on evaluation using a python script embedded within the problem.
</p>
  <p><span onclick="unhide('calculations_1')" style="cursor: pointer; color: blue; text-decoration: underline;">Click here</span> if you'd like to see the code used to create this kind of problem, which you can paste into your own course, and edit if needed.</p>
  <p>
    <div class="hidden" id="calculations_1" style="padding-bottom: 1em;">
      <pre>
        <code>
&lt;script type="loncapa/python"&gt;

def test_add(expect, ans):
    try:
        a1=int(ans[0])
        a2=int(ans[1])
        return (a1+a2) == int(expect)
    except ValueError:
        return False

def test_add_to_ten(expect, ans):
    return test_add(10, ans)

&lt;/script&gt;

&lt;p&gt;Enter two integers which sum to 10: &lt;/p&gt;
&lt;customresponse cfn="test_add_to_ten"&gt;
        &lt;textline size="40" correct_answer="3" label="Integer #1"/&gt;&lt;br/&gt;
        &lt;textline size="40" correct_answer="7" label="Integer #2"/&gt;
&lt;/customresponse&gt;

&lt;p&gt;Enter two integers which sum to 20: &lt;/p&gt;
&lt;customresponse cfn="test_add" expect="20"&gt;
        &lt;textline size="40" correct_answer="11" label="Integer #1"/&gt;&lt;br/&gt;
        &lt;textline size="40" correct_answer="9" label="Integer #2"/&gt;
&lt;/customresponse&gt;</code>
      </pre>
      <p>
        <span onclick="hide('calculations_1')" style="cursor: pointer; color: blue; text-decoration: underline;">Click here to hide this code.</span>
      </p>
    </div>
  </p>
  <script type="loncapa/python">

def test_add(expect, ans):
    try:
        a1=int(ans[0])
        a2=int(ans[1])
        return (a1+a2) == int(expect)
    except ValueError:
        return False

def test_add_to_ten(expect, ans):
    return test_add(10, ans)

</script>
  <p>Enter two integers which sum to 10: </p>
  <customresponse cfn="test_add_to_ten">
    <textline correct_answer="3" label="Integer #1" size="40"/>
    <br/>
    <textline correct_answer="7" label="Integer #2" size="40"/>
  </customresponse>
  <p>Enter two integers which sum to 20: </p>
  <customresponse cfn="test_add" expect="20">
    <textline correct_answer="11" label="Integer #1" size="40"/>
    <br/>
    <textline correct_answer="9" label="Integer #2" size="40"/>
  </customresponse>
  <solution>
    <div class="detailed-solution">
      <p>Explanation</p>
      <p>Any set of integers on the line \(y = 10 - x\) and \(y = 20 - x\) satisfy these constraints.</p>
      <img src="/static/images/simple_graph.png"/>
    </div>
  </solution>
</problem>
