<problem url_name="Problem_with_Adaptive_Hint_0" display_name="Problem with Adaptive Hint" markdown="null">
  <text>
    <p>
      <h4>Problem With Adaptive Hint</h4>
    </p>
    <p>
This problem demonstrates a question with hints, based on using the <tt class="tt">hintfn</tt> method. </p>
    <p>To create a problem like this of your own, you can use the following code as your starting point (please <span onclick="unhide('calculations_1')" style="cursor: pointer; color: blue; text-decoration: underline;">click on this link</span> to expose the code).  Note that a hint will only display if an answer you have expressly called out has been submitted.</p>
    <p>
      <div class="hidden" id="calculations_1" style="padding-bottom: 1em;">
        <pre>
          <code> &lt;script type="text/python" system_path="python_lib"&gt;
def test_str(expect, ans):
  print expect, ans
  ans = ans.strip("'")
  ans = ans.strip('"')
  return expect == ans.lower()

def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid]).lower()
  print 'hint_fn called, ans=', ans
  hint = ''
  if 'java' in ans:
     hint = 'that is only good for drinking'
  elif 'perl' in ans:
     hint = 'not that rich'
  elif 'pascal' in ans:
     hint = 'that is a beatnick language'
  elif 'fortran' in ans:
     hint = 'those were the good days'
  elif 'clu' in ans:
     hint = 'you must be invariant'
  if hint:
    hint = "&amp;lt;font color='blue'&amp;gt;Hint: {0}&amp;lt;/font&amp;gt;".format(hint)
    new_cmap.set_hint_and_mode(aid,hint,'always')
&lt;/script&gt;
    &lt;p&gt;
What is the best programming language that exists today? You may enter your answer in upper or lower case, with or without quotes. &lt;/p&gt;
    &lt;p&gt;
      &lt;customresponse cfn="test_str" expect="python"&gt;
        &lt;textline correct_answer="python" label="What is the best programming language that exists today?"/&gt;
        &lt;hintgroup hintfn="hint_fn"/&gt;
      &lt;/customresponse&gt;</code>
        </pre>
        <p>
          <span onclick="hide('calculations_1')" style="cursor: pointer; color: blue; text-decoration: underline;">Click here to hide this code.</span>
        </p>
      </div>
    </p>
    <script system_path="python_lib" type="text/python">
def test_str(expect, ans):
  print expect, ans
  ans = ans.strip("'")
  ans = ans.strip('"')
  return expect == ans.lower()

def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid]).lower()
  print 'hint_fn called, ans=', ans
  hint = ''
  if 'java' in ans:
     hint = 'that is only good for drinking'
  elif 'perl' in ans:
     hint = 'not that rich'
  elif 'pascal' in ans:
     hint = 'that is a beatnick language'
  elif 'fortran' in ans:
     hint = 'those were the good days'
  elif 'clu' in ans:
     hint = 'you must be invariant'
  if hint:
    hint = "&lt;font color='blue'&gt;Hint: {0}&lt;/font&gt;".format(hint)
    new_cmap.set_hint_and_mode(aid,hint,'always')
</script>
    <p>
What is the best programming language that exists today? You may enter your answer in upper or lower case, with or without quotes. </p>
    <p>
      <customresponse cfn="test_str" expect="python">
        <textline correct_answer="python" label="What is the best programming language that exists today?"/>
        <hintgroup hintfn="hint_fn"/>
      </customresponse>
    </p>
  </text>
</problem>
